# 19CSE367-DIP 
## Programming Task 3

##### In this assigment you will work with basic image read and write operations using opencv and scikit image libraries

First  we will start with open cv. Run the cell below to check if open cv is installed in your machine. if not the cell will automatically download the package and install it. 
In this case you need to restart the kernel.

```python
try:
    import cv2
except:
    !pip install opencv-python
```

```python
import cv2
```

<!-- grade_id: cell-20f68321026ce335, points: 0, grade: false, solution: true -->
```python
#Read the image "monalisa.jpg" as color image and save it in variable img_ocv_clr
### BEGIN SOLUTION
img_ocv_clr=cv2.imread('monalisa.jpg')
### END SOLUTION
```

<!-- grade_id: cell-f3075d5bb5e1c596, points: 1, grade: true, solution: false -->
```python
""" Check if the proper image was read for the task"""
assert img_ocv_clr.shape==(256, 256, 3)
assert img_ocv_clr.size== 196608
assert img_ocv_clr.ndim==3
assert img_ocv_clr.dtype == 'uint8'
### BEGIN HIDDEN TESTS
assert img_ocv_clr[255][250][0]==63
### END HIDDEN TESTS
```

<!-- grade_id: cell-e3c661a70f87deca, points: 0, grade: false, solution: true -->
```python
#Read the image "monalisa.jpg" as grayscale image and save it in variable img_ocv_gray
### BEGIN SOLUTION
img_ocv_gray=cv2.imread('monalisa.jpg',0)
### END SOLUTION
```

<!-- grade_id: cell-75370523605087b7, points: 1, grade: true, solution: false -->
```python
""" Check if the proper image was read for the task"""
assert img_ocv_gray.shape==(256, 256)
assert img_ocv_gray.size== 65536
assert img_ocv_gray.ndim==2
assert img_ocv_gray.dtype == 'uint8'
### BEGIN HIDDEN TESTS
assert img_ocv_gray[55][100]==179
### END HIDDEN TESTS
```

Lets move on to Scikit image library.Run the cell below to check if scikit image is installed in your machine. if not the cell will automatically download the package and install it. 

```python
try:
    import skimage
except:
    !pip install scikit-image
    
```

```python
# import the scikit image lib or specific module used for reading an image
from skimage import io
```

<!-- grade_id: cell-fd4d3f40b78035b1, points: 0, grade: false, solution: true -->
```python
#Read the image "monalisa.jpg" as color image and save it in variable img_sk_clr
### BEGIN SOLUTION
img_sk_clr=io.imread("monalisa.jpg",)
### END SOLUTION
```

<!-- grade_id: cell-3b18662c28e35a44, points: 2, grade: true, solution: false -->
```python
""" Check if the proper image was read for the task"""
assert img_sk_clr.shape==(256, 256, 3)
assert img_sk_clr.size== 196608
assert img_sk_clr.ndim==3
assert img_sk_clr.dtype == 'uint8'
### BEGIN HIDDEN TESTS
assert img_sk_clr[255][250][0]==60
### END HIDDEN TESTS
```

<!-- grade_id: cell-3839347599835791, points: 0, grade: false, solution: true -->
```python
#Read the image "monalisa.jpg" as grayscale image and save it in variable img_sk_gray
#hint : Look into the documentation for imread from scikit image, there is an argument which enables reading images as grayscale
### BEGIN SOLUTION
img_sk_gray=io.imread('monalisa.jpg',as_gray=True)
### END SOLUTION
```

<!-- grade_id: cell-8008023df1ba6439, points: 2, grade: true, solution: false -->
```python
""" Check if the proper image was read for the task"""
assert img_sk_gray.shape==(256, 256)
assert img_sk_gray.size== 65536
assert img_sk_gray.ndim==2
assert img_sk_gray.dtype == 'float64'
### BEGIN HIDDEN TESTS
assert img_sk_gray[255][6]==0.19900039215686274
### END HIDDEN TESTS
```

### Color space conversion

```python
# run this cell to install imagecodecs else you may get  error while working with tif file. 

pip install imagecodecs
#!pip install imagecodecs
```

```python
import cv2
from skimage import io

img_ocv_clr=cv2.imread('img1.tif')# BGR
img_sk_clr=io.imread("img1.tif") # RGB
```

<!-- grade_id: cell-5a55d6c11afabe50, points: 0, grade: false, solution: true -->
```python
#Convert the imgage loaded using scikit image(img_sk_clr) to BGR color space 
# save the converted image in  img_sk_clr_BGR variable
# note that images loaded using scikit is in RGB format whereas images loaded using open cv is BGR
# Note - use cvtColor function for this task
### BEGIN SOLUTION
img_sk_clr_BGR=cv2.cvtColor(img_sk_clr,cv2.COLOR_RGB2BGR)
### END SOLUTION
```

<!-- grade_id: cell-aba23d978c2cace8, points: 2, grade: true, solution: false -->
```python
""" Check if the right conversion was done or not"""
assert img_sk_clr_BGR[0][0][0] != img_sk_clr[0][0][0]
assert img_sk_clr_BGR[0][0][0]==img_ocv_clr[0][0][0]
### BEGIN HIDDEN TESTS
assert (img_sk_clr_BGR == img_ocv_clr).all()
### END HIDDEN TESTS
```

## Writing an image into file

<!-- grade_id: cell-7df91cc072ca0301, points: 0, grade: false, solution: true -->
```python
import cv2
# load the image monalis.jpg as gray scale image using open cv and save it in variable img_ocv_gs

### BEGIN SOLUTION
img_ocv_gs=cv2.imread('monalisa.jpg',0)
### END SOLUTION
```

```python
# this operation is for generating the negative of a 8 bit image.
neg_image=255-img_ocv_gs
```

<!-- grade_id: cell-e2cc5c52d6379151, points: 0, grade: false, solution: true -->
```python
# Using  opencv write the gray scale version of monalisa to file with name "gs_monalisa.jpg"
### BEGIN SOLUTION
cv2.imwrite("gs_monalisa.jpg",img_ocv_gs)
### END SOLUTION
```

Output:
```
True
```

<!-- grade_id: cell-79a6e8d0e6b16046, points: 0, grade: false, solution: true -->
```python
# Using  opencv write the negative image of gray scale monalisa to file with name "neg_monalisa.jpg"
### BEGIN SOLUTION
cv2.imwrite("neg_monalisa.jpg",neg_image)
### END SOLUTION
```

Output:
```
True
```

<!-- grade_id: cell-8a8d4166e93f76c2, points: 2, grade: true, solution: false -->
```python
""" Check if the writing of image is properly done"""
img_test=cv2.imread('monalisa.jpg')
img_test_gs=cv2.imread('gs_monalisa.jpg',0)
img_test_neg=cv2.imread('neg_monalisa.jpg',0)

assert img_test.shape != img_test_gs.shape
assert img_test_gs.shape==(256,256)
### BEGIN HIDDEN TESTS
assert img_test_neg[11][15]== 255-img_test_gs[11][15]
### END HIDDEN TESTS
```